Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ben Bitdiddle"
COLLABORATORS = "Alyssa P. Hacker"
```

---

## Part A

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    if n < 1:
        raise ValueError
    s = []
    for i in range(n):
        s.append(i**2)
    return s
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

Output:
```
[0, 1, 4, 9, 16, 25, 36, 49, 64, 81]
```

<!-- grade_id: correct_squares, points: 1.0, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from nose.tools import assert_equal
assert_equal(squares(1), [1])
assert_equal(squares(2), [1, 4])
assert_equal(squares(10), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100])
assert_equal(squares(11), [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121])
```

Output:
```
/opt/python/3.4.2/lib/python3.4/unittest/case.py in assertSequenceEqual(self, seq1, seq2, msg, seq_type)
    975         standardMsg = self._truncateMessage(standardMsg, diffMsg)
    976         msg = self._formatMessage(msg, standardMsg)
--> 977         self.fail(msg)
    978 
    979     def _truncateMessage(self, message, diff):

/opt/python/3.4.2/lib/python3.4/unittest/case.py in fail(self, msg)
    640     def fail(self, msg=None):
    641         """Fail immediately, with the given message."""
--> 642         raise self.failureException(msg)
    643 
    644     def assertFalse(self, expr, msg=None):

AssertionError: Lists differ: [0] != [1]

First differing element 0:
0
1

- [0]
+ [1]
```

<!-- grade_id: squares_invalid_input, points: 1.0, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
from nose.tools import assert_raises
assert_raises(ValueError, squares, 0)
assert_raises(ValueError, squares, -4)
```

---

## Part B

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    total = 0
    s = squares(n)
    for i in range(len(s)):
        total += s[i]
    return total
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

Output:
```
285
```

<!-- grade_id: correct_sum_of_squares, points: 1.0, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
assert_equal(sum_of_squares(1), 1)
assert_equal(sum_of_squares(2), 5)
assert_equal(sum_of_squares(10), 385)
assert_equal(sum_of_squares(11), 506)
```

Output:
```
/opt/python/3.4.2/lib/python3.4/unittest/case.py in assertEqual(self, first, second, msg)
    795         """
    796         assertion_func = self._getAssertEqualityFunc(first, second)
--> 797         assertion_func(first, second, msg=msg)
    798 
    799     def assertNotEqual(self, first, second, msg=None):

/opt/python/3.4.2/lib/python3.4/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    788             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    789             msg = self._formatMessage(msg, standardMsg)
--> 790             raise self.failureException(msg)
    791 
    792     def assertEqual(self, first, second, msg=None):

AssertionError: 0 != 1
```

<!-- grade_id: sum_of_squares_uses_squares, points: 1.0, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""
orig_squares = squares
del squares
try:
    assert_raises(NameError, sum_of_squares, 1)
except AssertionError:
    raise AssertionError("sum_of_squares does not use squares")
finally:
    squares = orig_squares
```

---
## Part C

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1.0, grade: true, solution: true -->
$\sum_{i=0}^n i^2$